<a href="https://colab.research.google.com/github/ferngonzalezp/deep_learning_lab/blob/main/2022/TP2/TP2_Overview_on_Numpy_and_Pandas.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# 1. A tour through Numpy

*material adapted from [Scientific Computing in Python: Introduction to NumPy and Matplotlib](https://sebastianraschka.com/blog/2020/numpy-intro.html).*
________________________________________________________________________________


This section offers a quick tour of the NumPy library for working with multi-dimensional arrays in Python. NumPy (short for Numerical Python) was created in 2005 by merging Numarray into Numeric. Since then, the open source NumPy library has evolved into an essential library for scientific computing in Python. It has become a building block of many other scientific libraries, such as SciPy, Scikit-learn, Pandas, and others. What makes NumPy so incredibly attractive to the scientific community is that it provides a convenient Python interface for working with multi-dimensional array data structures efficiently; the NumPy array data structure is also called ndarray, which is short for n-dimensional array.

In addition to being mostly implemented in C and using Python as a “glue language,” the main reason why NumPy is so efficient for numerical computations is that NumPy arrays use contiguous blocks of memory that can be efficiently cached by the CPU. In contrast, Python lists are arrays of pointers to objects in random locations in memory, which cannot be easily cached and come with a more expensive memory-look-up. However, the computational efficiency and low-memory footprint come at a cost: NumPy arrays have a fixed size and are homogeneous, which means that all elements must have the same type. Homogenous ndarray objects have the advantage that NumPy can carry out operations using efficient C code and avoid expensive type checks and other overheads of the Python API. While adding and removing elements from the end of a Python list is very efficient, altering the size of a NumPy array is very expensive since it requires to create a new array and carry over the contents of the old array that we want to expand or shrink.

Besides being more efficient for numerical computations than native Python code, NumPy can also be more elegant and readable due to vectorized operations and broadcasting, which are features that we will explore in this article.

Today, NumPy forms the basis of the scientific Python computing ecosystem.

## Motivation: NumPy is fast!

Here is some motivation before we discuss further details, highlighting why learning about and using NumPy is useful. We take a look at a speed comparison with regular Python code. In particular, we are computing a vector dot product in Python (using lists) and compare it with  NumPy's dot-product function. Mathematically, the dot product between two vectors $\mathbf{x}$ and $\mathbf{w}$ can be written as follows:

$$
z = \sum_i x_i w_i = x_1 \times w_1 + x_2 \times w_2 + ... + x_n \times w_n = \mathbf{x}^\top \mathbf{w} 
$$

First, the Python implementation using a for-loop:

```python
def python_forloop_list_approach(x, w):
    z = 0.
    for i in range(len(x)):
        z += x[i] * w[i]
    return z


a = [1., 2., 3.]
b = [4., 5., 6.]

print(python_forloop_list_approach(a, b))
```

Let us compute the runtime for two larger (1000-element) vectors using IPython's `%timeit` magic function:

```python
large_a = list(range(1000))
large_b = list(range(1000))


%timeit python_forloop_list_approach(large_a, large_b)
```

Next, we use the `dot` function/method implemented in NumPy to compute the dot product between two vectors and run `%timeit` afterwards:

```python
import numpy as np


def numpy_dotproduct_approach(x, w):
    # same as np.dot(x, w)
    # and same as x @ w
    return x.dot(w)
    

a = np.array([1., 2., 3.])
b = np.array([4., 5., 6.])

print(numpy_dotproduct_approach(a, b))
```

```python
large_a = np.arange(1000)
large_b = np.arange(1000)

%timeit numpy_dotproduct_approach(large_a, large_b)
```

As we can see, replacing the for-loop with NumPy's dot function makes the computation of the vector dot product approximately 100 times faster.

## N-dimensional Arrays

NumPy is built around [`ndarrays`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.ndarray.html) objects, which are high-performance multi-dimensional array data structures. Intuitively, we can think of a one-dimensional NumPy array as a data structure to represent a vector of elements -- you may think of it as a fixed-size Python list where all elements share the same type. Similarly, we can think of a two-dimensional array as a data structure to represent a matrix or a Python list of lists. While NumPy arrays can have up to 32 dimensions if it was compiled without alterations to the source code, we will focus on lower-dimensional arrays for the purpose of illustration in this introduction.

Now, let us get started with NumPy by calling the `array` function to create a two-dimensional NumPy array, consisting of two rows and three columns, from a list of lists:

```python
a = [1., 2., 3.]
np.array(a)
```

```python
lst = [[1, 2, 3], 
       [4, 5, 6]]
ary2d = np.array(lst)
ary2d

# rows x columns
```

<img src="https://drive.google.com/uc?id=1p9U2qao0KqP84fcfa7VEtBV3nXXjuBco" alt="Drawing" style="width: 300px;"/>

By default, NumPy infers the type of the array upon construction. Since we passed Python integers to the array, the `ndarray` object `ary2d` should be of type `int64` on a 64-bit machine, which we can confirm by accessing the `dtype` attribute:

```python
ary2d.dtype
```

If we want to construct NumPy arrays of different types, we can pass an argument to the `dtype` parameter of the `array` function, for example `np.int32`, to create 32-bit arrays. For a full list of supported data types, please refer to the official [NumPy documentation](https://docs.scipy.org/doc/numpy/user/basics.types.html). Once an array has been constructed, we can downcast or recast its type via the `astype` method as shown in the following examples:

```python
int32_ary = ary2d.astype(np.int32)
int32_ary
```

```python
float32_ary = ary2d.astype(np.float32)
float32_ary
```

```python
float32_ary.dtype
```

The code snippet above returned `8`, which means that each element in the array (remember that `ndarray`s are homogeneous) takes up 8 bytes in memory. This result makes sense since the array `ary2d` has type `int64` (64-bit integer), which we determined earlier, and 8 bits equals 1 byte. (Note that `'int64'` is just a shorthand for `np.int64`.)

To return the number of elements in an array, we can use the `size` attribute, as shown below:

And the number of dimensions of our array (Intuitively, you may think of *dimensions* as the *rank* of a tensor) can be obtained via the `ndim` attribute:

```python
ary2d
```

```python
ary2d.ndim
```

If we are interested in the number of elements along each array dimension (in the context of NumPy arrays, we may also refer to them as *axes*), we can access the `shape` attribute as shown below:

```python
len(ary2d.shape)
```

The `shape` is always a tuple; in the code example above, the two-dimensional `ary` object has two *rows* and *three* columns, `(2, 3)`, if we think of it as a matrix representation.

Conversely, the `shape` (an object of type `tuple`) of a one-dimensional array only contains a single value:

```python
np.array([1., 2., 3.]).shape
```

---

## NumPy Array Construction and Indexing

---

### Array Construction Routines

This section provides a non-comprehensive list of array construction functions. Simple yet useful functions exist to construct arrays containing ones or zeros:

```python
np.ones((3, 4), dtype=np.int)
```

```python
np.zeros((3, 3))
```

We can use these functions to create arrays with arbitrary values, e.g., we can create an array containing the values 99 as follows:

```python
np.zeros((3, 3)) + 99
```

Creating arrays of ones or zeros can also be useful as placeholder arrays, in cases where we do not want to use the initial values for computations but want to fill it with other values right away. If we do not need the initial values (for instance, `'0.'` or `'1.'`), there is also `numpy.empty`, which follows the same syntax as `numpy.ones` and `np.zeros`. However, instead of filling the array with a particular value, the `empty` function creates the array with non-sensical values from memory. We can think of `zeros` as a function that creates the array via `empty` and then sets all its values to `0.` -- in practice, a difference in speed is not noticeable, though.  

NumPy also comes with functions to create identity matrices and diagonal matrices as `ndarrays` that can be useful in the context of linear algebra -- a topic that we will explore later in this article. 

```python
np.eye(3)
```

```python
np.diag((1, 2, 3))
```

Lastly, I want to mention two very useful functions for creating sequences of numbers within a specified range, namely, `arange` and `linspace`. NumPy's `arange` function follows the same syntax as Python's `range` objects: If two arguments are provided, the first argument represents the start value and the second value defines the stop value of a half-open interval:

```python
np.arange(4, 10)
```

Notice that `arange` also performs type inference similar to the `array` function. If we only provide a single function argument, the range object treats this number as the endpoint of the interval and starts at 0:

```python
np.arange(5)
```

Similar to Python's `range`, a third argument can be provided to define the *step* (the default step size is 1). For example, we can obtain an array of all uneven values between one and ten as follows:

```python
np.arange(1., 11., 0.1)
```

The `linspace` function is especially useful if we want to create a particular number of evenly spaced values in a specified half-open interval:

```python
np.linspace(6., 15., num=10)
```

### Array Indexing

In this section, we will go over the basics of retrieving NumPy array elements via different indexing methods. Simple NumPy indexing and slicing works similar to Python lists, which we will demonstrate in the following code snippet, where we retrieve the first element of a one-dimensional array:

```python
ary = np.array([1, 2, 3])
ary[0]
```

Also, the same Python semantics apply to slicing operations. The following example shows how to fetch the first two elements in `ary`:

```python
ary[0:3] # equivalent to ary[0:2]
```

If we work with arrays that have more than one dimension or axis, we separate our indexing or slicing operations by commas as shown in the series of examples below:

```python
ary = np.array([[1, 2, 3],
                [4, 5, 6]])

ary[0, -2] # first row, second from last element
```

```python
ary[-1, -1] # lower right
```

```python
ary[1, 1] # first row, second column
```

<img src="https://drive.google.com/uc?id=1cfCht7TEsU9NTga9s4GO6jmNwI0GkKIH" alt="Drawing" style="width: 300px;"/>

```python
ary[:, 0] # entire first column
```

---

## NumPy Array Math and Universal Functions

---

### Array Math and Universal Functions

In the previous sections, you learned how to create NumPy arrays and how to access different elements in an array. It is about time that we introduce one of the core features of NumPy that makes working with `ndarray` so efficient and convenient: vectorization. While we typically use for-loops if we want to perform arithmetic operations on sequence-like objects, NumPy provides vectorized wrappers for performing element-wise operations implicitly via so-called *ufuncs* -- "*ufuncs*" is short for universal functions.

As of this writing, there are more than 60 ufuncs available in NumPy; ufuncs are implemented in compiled C code and very fast and efficient compared to vanilla Python. In this section, we will take a look at the most commonly used ufuncs, and I recommend you to check out the [official documentation](https://docs.scipy.org/doc/numpy/reference/ufuncs.html#available-ufuncs) for a complete list.

To provide an example of a simple ufunc for element-wise addition, consider the following example, where we add a scalar (here: 1) to each element in a nested Python list:

```python
lst = [[1, 2, 3], 
       [4, 5, 6]] # 2d array

for row_idx, row_val in enumerate(lst):
    for col_idx, col_val in enumerate(row_val):
        lst[row_idx][col_idx] += 1
lst
```

This for-loop approach is very verbose, and we could achieve the same goal more elegantly using list comprehensions:

```python
lst = [[1, 2, 3], [4, 5, 6]]
[[cell + 1 for cell in row] for row in lst]
```

We can accomplish the same using NumPy's ufunc for element-wise scalar addition as shown below:

```python
ary = np.array([[1, 2, 3], [4, 5, 6]])
ary = np.add(ary, 1) # binary ufunc
ary
```

The ufuncs for basic arithmetic operations are `add`, `subtract`, `divide`, `multiply`, `power`, and `exp` (exponential). However, NumPy uses operator overloading so that we can use mathematical operators (`+`, `-`, `/`, `*`, and `**`) directly:

```python
np.add(ary, 1)
```

```python
ary + 1
```

```python
np.power(ary, 2)
```

```python
ary**2
```

Above, we have seen examples of *binary* ufuncs, which are ufuncs that take two arguments as an input. In addition, NumPy implements several useful *unary* ufuncs, such as `log` (natural logarithm), `log10` (base-10 logarithm), and `sqrt` (square root):



```python
np.sqrt(ary)
```

Often, we want to compute the sum or product of array element along a given axis. For this purpose, we can use a ufunc's `reduce` operation. By default, `reduce` applies an operation along the first axis (`axis=0`). In the case of a two-dimensional array, we can think of the first axis as the rows of a matrix. Thus, adding up elements along rows yields the column sums of that matrix as shown below:

```python
ary = np.array([[1, 2, 3], 
                [4, 5, 6]]) # rolling over the 1st axis, axis 0

np.add.reduce(ary, axis=0)
```

To compute the row sums of the array above, we can specify `axis=1`:

```python
np.add.reduce(ary, axis=1) # row sums
```

While it can be more intuitive to use `reduce` as a more general operation, NumPy also provides shorthands for specific operations such as `product` and `sum`. For example, `sum(axis=0)` is equivalent to `add.reduce`:

```python
ary.sum(axis=0) # column sums
```

```python
ary.sum(axis=1) # row sums
```

<img src="https://drive.google.com/uc?id=1dyAwvgJbiHOSD2pJkIoQrEI-Myn9mGLx" alt="Drawing" style="width: 300px;"/>

As a word of caution, keep in mind that `product` and `sum` both compute the product or sum of the entire array if we do not specify an axis:

```python
ary.sum()
```

Other useful unary ufuncs are:
    
- `np.mean` (computes arithmetic mean or average)
- `np.std` (computes the standard deviation)
- `np.var` (computes variance)
- `np.sort` (sorts an array)
- `np.argsort` (returns indices that would sort an array)
- `np.min` (returns the minimum value of an array)
- `np.max` (returns the maximum value of an array)
- `np.argmin` (returns the index of the minimum value)
- `np.argmax` (returns the index of the maximum value)
- `np.array_equal` (checks if two arrays have the same shape and elements)

```python
ary[:, :2] # first two columns
```

---

## NumPy Broadcasting

---

### Broadcasting

A topic we glanced over in the previous section is broadcasting. Broadcasting allows us to perform vectorized operations between two arrays even if their dimensions do not match by creating implicit multidimensional grids. You already learned about ufuncs in the previous section where we performed element-wise addition between a scalar and a multidimensional array, which is just one example of broadcasting. 

<img src="https://drive.google.com/uc?id=1O7AICJkfUeiDdEh9mR4EFu5UsvCqqXm_" alt="Drawing" style="width: 300px;"/>

Naturally, we can also perform element-wise operations between arrays of equal dimensions:

In contrast to what we are used from linear algebra, we can also add arrays of different shapes. In the example above, we will add a one-dimensional to a two-dimensional array, where NumPy creates an implicit multidimensional grid from the one-dimensional array `ary1`:

```python
ary = np.array([1, 2, 3])
ary + 1
```

```python
ary + np.array([1, 1, 1])
```

<img src="https://drive.google.com/uc?id=10S5E0QObu90rFIo3BoMDFLA-8S6kxhh_" alt="Drawing" style="width: 300px;"/>

```python
ary2 = np.array([[4, 5, 6], 
                 [7, 8, 9]])

ary2 + ary
```

---

## Reshaping NumPy Arrays

---

### Reshaping Arrays

In practice, we often run into situations where existing arrays do not have the *right* shape to perform certain computations. As you might remember from the beginning of this article, the size of NumPy arrays is fixed. Fortunately, this does not mean that we have to create new arrays and copy values from the old array to the new one if we want arrays of different shapes -- the size is fixed, but the shape is not. NumPy provides a `reshape` methods that allow us to obtain a view of an array with a different shape. 

For example, we can reshape a one-dimensional array into a two-dimensional one using `reshape` as follows:

```python
ary1d = np.array([1, 2, 3, 4, 5, 6])
ary2d_view = ary1d.reshape(2, 3)
ary2d_view
```

The `True` value returned from `np.may_share_memory` indicates that the reshape operation returns a memory view, not a copy:

```python
np.may_share_memory(ary2d_view, ary1d)
```

While we need to specify the desired elements along each axis, we need to make sure that the reshaped array has the same number of elements as the original one. However, we do not need to specify the number elements in each axis; NumPy is smart enough to figure out how many elements to put along an axis if only one axis is unspecified (by using the placeholder `-1`):

```python
ary1d.reshape(-1, 2)
```

We can, of course, also use `reshape` to flatten an array:

```python
ary = np.array([[[1, 2, 3],
                 [4, 5, 6]]])

ary.reshape(-1)
```

Other methods for flattening arrays exist, namely `flatten`, which creates a copy of the array, and `ravel`, which creates a memory view like reshape:

```python
ary.flatten()
```

```python
ary.ravel()
```

Sometimes, we are interested in merging different arrays. Unfortunately, there is no efficient way to do this without creating a new array, since NumPy arrays have a fixed size. While combining arrays should be avoided if possible -- for reasons of computational efficiency -- it is sometimes necessary. To combine two or more array objects, we can use NumPy's `concatenate` function as shown in the following examples:

```python
ary = np.array([1, 2, 3])

# stack along the first axis
np.concatenate((ary, ary)) 
```

```python
ary = np.array([[1, 2, 3]])

# stack along the first axis (here: rows)
np.concatenate((ary, ary), axis=0)
```

```python
# stack along the second axis (here: column)
np.concatenate((ary, ary), axis=1)
```

---

## Random Number Generators

---

### Random Number Generators

In machine learning and deep learning, we often have to generate arrays of random numbers -- for example, the initial values of our model parameters before optimization. NumPy has a `random` subpackage to create random numbers and samples from a variety of distributions conveniently. Again, I encourage you to browse through the more comprehensive [numpy.random documentation](https://docs.scipy.org/doc/numpy-1.14.0/reference/routines.random.html) for a complete list of functions for random sampling.

To provide a brief overview of the pseudo-random number generators that we will use most commonly, let's start with drawing a random sample from a uniform distribution:

```python
np.random.seed(123)

np.random.rand(3)
```

In the code snippet above, we first seeded NumPy's random number generator. Then, we drew three random samples from a uniform distribution via `random.rand` in the half-open interval [0, 1). I highly recommend the seeding step in practical applications as well as in research projects, since it ensures that our results are reproducible. If we run our code sequentially -- for example, if we execute a Python script -- it should be sufficient to seed the random number generator only once at the beginning to enforce reproducible outcomes between different runs. However, it is often useful to create separate `RandomState` objects for various parts of our code, so that we can test methods of functions reliably in unit tests. Working with multiple, separate `RandomState` objects can also be useful if we run our code in non-sequential order -- for example if we are experimenting with our code in interactive sessions or Jupyter Notebook environments. 

The example below shows how we can use a `RandomState` object to create the same results that we obtained via `np.random.rand` in the previous code snippet:

```python
rng2 = np.random.RandomState(seed=531)
rng2.rand(3)
```

Also, the NumPy developer community developed new random number generation method in recent versions of NumPy. For more details, please see the [new random `Generator` documentation](https://numpy.org/doc/stable/reference/random/generator.html#numpy.random.default_rng):

```python
rng2 = np.random.default_rng(seed=123)
rng2.random(3)
```

---

## Linear Algebra with NumPy

---

### Linear Algebra with NumPy Arrays

Intuitively, we can think of one-dimensional NumPy arrays as data structures that represent row vectors:

```python
row_vector = np.array([1, 2, 3])
row_vector
```

Similarly, we can use two-dimensional arrays to create column vectors:

```python
column_vector = np.array([1, 2, 3]).reshape(-1, 1)
column_vector
```

Instead of reshaping a one-dimensional array into a two-dimensional one, we can simply add a new axis as shown below:

```python
row_vector[:, np.newaxis]
```

Note that in this context, `np.newaxis` behaves like `None`:

```python
row_vector[:, None]
```

All three approaches listed above, using `reshape(-1, 1)`, `np.newaxis`, or `None` yield the same results -- all three approaches create views not copies of the `row_vector` array.

We can think of a column vector as a matrix consisting only of one column. To perform matrix multiplication between matrices, we learned that number of columns of the left matrix must match the number of rows of the matrix to the right. In NumPy, we can perform matrix multiplication via the `matmul` function:

```python
matrix = np.array([[1, 2, 3], 
                   [4, 5, 6]])
```

```python
np.matmul(matrix, column_vector)
```

<img src="https://drive.google.com/uc?id=1yG9ixNkz52WBApp_GHBi7sTKBAAEG3TK"  alt="Drawing" style="width: 300px;"/>

However, if we are working with matrices and vectors, NumPy can be quite forgiving if the dimensions of matrices and one-dimensional arrays do not match exactly -- thanks to broadcasting. The following example yields the same result as the matrix-column vector multiplication, except that it returns a one-dimensional array instead of a two-dimensional one:

```python
np.matmul(matrix, row_vector)
```

Similarly, we can compute the dot-product between two vectors (here: the vector norm)

```python
np.matmul(row_vector, row_vector)
```

NumPy has a special `dot` function that behaves similar to `matmul` on pairs of one- or two-dimensional arrays -- its underlying implementation is different though, and one or the other can be slightly faster on specific machines and versions of [BLAS](https://en.wikipedia.org/wiki/Basic_Linear_Algebra_Subprograms):

```python
np.dot(matrix, row_vector)
```

Note that an even more convenient way for executing `np.dot` is using the `@` symbol with NumPy arrays:

```python
matrix @ row_vector
```

Similar to the examples above we can use `matmul` or `dot` to multiply two matrices (here: two-dimensional arrays). In this context, NumPy arrays have a handy `transpose` method to transpose matrices if necessary:

```python
matrix = np.array([[1, 2, 3], 
                   [4, 5, 6]])

matrix.transpose()
```

<img src="https://drive.google.com/uc?id=1wz4hORP4e01wjE-Tn1Txx1JQ7uanyFJ1"  alt="Drawing" style="width: 500px;"/>

```python
np.dot(matrix, matrix.transpose())
```

<img src="https://drive.google.com/uc?id=1yCzWMZ0ql3MLGuuq2FeOxQqXhDdB8fvE"  alt="Drawing" style="width: 500px;"/>

While `transpose` can be annoyingly verbose for implementing linear algebra operations -- think of [PEP8's](https://www.python.org/dev/peps/pep-0008/) *80 character per line* recommendation -- NumPy has a shorthand for that: `T`:

```python
matrix.T
```

While this section demonstrates some of the basic linear algebra operations carried out on NumPy arrays that we use in practice, you can find an additional function in the documentation of NumPy's submodule for linear algebra: [`numpy.linalg`](https://docs.scipy.org/doc/numpy/reference/routines.linalg.html). If you want to perform a particular linear algebra routine that is not implemented in NumPy, it is also worth consulting the [`scipy.linalg` documentation](https://docs.scipy.org/doc/scipy/reference/linalg.html) -- SciPy is a library for scientific computing built on top of NumPy.

---

I want to mention that there is also a special [`matrix`](https://docs.scipy.org/doc/numpy/reference/generated/numpy.matrix.html) type in NumPy. NumPy `matrix` objects are analogous to NumPy arrays but are restricted to two dimensions. Also, matrices define certain operations differently than arrays; for instance, the `*` operator performs matrix multiplication instead of element-wise multiplication. However, NumPy `matrix` is less popular in the science community compared to the more general array data structure. 

---

## SciPy

SciPy is another open-source library from Python's scientific computing stack. SciPy includes submodules for integration, optimization, and many other kinds of computations that are out of the scope of NumPy itself. We will not cover SciPy as a library here, since it can be more considered as an "add-on" library on top of NumPy. 

I recommend you to take a look at the SciPy documentation to get a brief overview of the different function that exists within this library: [https://docs.scipy.org/doc/scipy/reference/](https://docs.scipy.org/doc/scipy/reference/)

## Assignment 1: Write a function to efficiently compute the Frobenius norm



The Frobenius norm, sometimes also called the Euclidean norm (a term unfortunately also used for the vector L^2-norm), is matrix norm of an m×n matrix A defined as the square root of the sum of the absolute squares of its elements.

\begin{equation}
 ||A||_F = \sqrt{\sum_{i=1}^m \sum_{j=1}^n|a_{ij}|^2}
\end{equation}

First construct a function using loops:


<!-- grade_id: cell-04522f42afe3f1f0, points: 0, grade: false, solution: true -->
```python
#function of Frobenius norm using loops:


def fro(A):
# YOUR CODE HERE
raise NotImplementedError()
```

Now eliminate the loops for a more efficient function. Hint: try to think how you can do this by doing a single dot product.

<!-- grade_id: cell-f2b81001fdf23f4d, points: 0, grade: false, solution: true -->
```python
def fro_efficient(A):
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-77ae5207f65e8a98, points: 2, grade: true, solution: false -->
```python
#test it!
A = np.ones((100,100))
%time print(fro(A))
%time print(fro_efficient(A))
%time print(np.linalg.norm(A))
assert(fro_efficient(A) == np.linalg.norm(A))
assert(fro(A) == np.linalg.norm(A))
```

## Assingment 2: The Kronecker Product


In mathematics, the Kronecker product, sometimes denoted by ⊗, is an operation on two matrices of arbitrary size resulting in a block matrix. If A is an $m × n$ matrix and B is a $p × q$ matrix, then the Kronecker product A ⊗ B is the $pm × qn$ block matrix:

\begin{equation}
\mathbf {A} \otimes \mathbf {B} ={\begin{bmatrix}a_{11}\mathbf {B} &\cdots &a_{1n}\mathbf {B} \\\vdots &\ddots &\vdots \\a_{m1}\mathbf {B} &\cdots &a_{mn}\mathbf {B} \end{bmatrix}}
\end{equation}

The numpy API has integrated a function to calculate the Kronecker product called `numpy.kron` [see docs](https://numpy.org/doc/stable/reference/generated/numpy.kron.html). The assingment consist of writing a function for the Kronecker product. Hint: you can use the function `numpy.einsum`, to see how it works consult the [documentation](https://numpy.org/doc/stable/reference/generated/numpy.einsum.html).

```python
A = np.array(([1,-4,-7],[-2,3,3]))
B = np.array(([8,-9,-6,5],[1,-3,-4,7],[2,8,-8,3],[1,2,-5,-1]))
%time C = np.kron(A,B)
print(A)
print(B)
print(C)
```

<!-- grade_id: cell-60160d21d0a91ccd, points: 0, grade: false, solution: true -->
```python
#Write the function of the Kronecker product

def kronecker(A,B):
  # YOUR CODE HERE
  raise NotImplementedError()
```

<!-- grade_id: cell-1349f70cab3a7158, points: 2, grade: true, solution: false -->
```python
#Test it
%time C = np.kron(A,B)
%time C2 = kronecker(A,B) #By using einsum, this function can be faster than current numpy implementation
print(C2)
```

# 2. Introduction to Pandas

Read the guide [10 minutes to pandas](https://pandas.pydata.org/docs/user_guide/10min.html), then download the electricity dataset from [OpenML](https://www.openml.org/d/151). Use Pandas to explore the dataset.

```python
import pandas as pd
import os
import matplotlib.pyplot as plt
```

**Question 1. Load the electricity dataset from the .csv (name your variable ds)**

<!-- grade_id: cell-9fc6bd940ea75ab4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a769b6a3d87aea05, points: 1, grade: true, solution: false -->
```python
ds.head()
```

**Question 2. What is the mean and standard deviation of the New South Wales price?**

<!-- grade_id: cell-d2c5f77a61480a21, points: 0, grade: false, solution: true -->
```python
def mean_std_nswprice(ds):
  # YOUR CODE HERE
  raise NotImplementedError()
  return mean, std
```

<!-- grade_id: cell-22d68bf4f6f9728d, points: 1, grade: true, solution: false -->
```python
mean, std = mean_std_nswprice(ds)
print(mean)
print(std)
```

**Question 3. Plot the histogram of the demand in Victoria with 10 bins**

<!-- grade_id: cell-c3932bf2ed5ad019, points: 0, grade: false, solution: true -->
```python
def plot_hist(ds):
  # YOUR CODE HERE
  raise NotImplementedError()
  return ax
```

<!-- grade_id: cell-a06ec003edfdae22, points: 1, grade: true, solution: false -->
```python
plt.figure()
ax = plot_hist(ds)
```

**Question 4. What proportion of the time , the price increased?**

<!-- grade_id: cell-e5d3231df18df8b9, points: 0, grade: false, solution: true -->
```python
def get_percent_up(ds):
  # YOUR CODE HERE
  raise NotImplementedError()
```

<!-- grade_id: cell-60734699a5e95c54, points: 1, grade: true, solution: false -->
```python
print(get_percent_up(ds))
```

**Question 5. Draw in a single plot the curve of the prices of New South Wales on fridays for the last quarter of the period**

<!-- grade_id: cell-eff2c4851a5e35e6, points: 0, grade: false, solution: true -->
```python
def plot_nswprices(ds):
  # YOUR CODE HERE
  raise NotImplementedError()
  return ax
```

<!-- grade_id: cell-1b6a3c7c960e1c95, points: 1, grade: true, solution: false -->
```python
plt.figure()
ax = plot_nswprices(ds)
```

**Question 6. Search the documentation and plot the probability density function of the demand of New Soth Wales using the kernel density estimate.**

<!-- grade_id: cell-8dc1b40a855309bf, points: 0, grade: false, solution: true -->
```python
def plot_nswdemand_pdf(ds):
  # YOUR CODE HERE
  raise NotImplementedError()
  return ax
```

<!-- grade_id: cell-e01bd3a17e301b87, points: 1, grade: true, solution: false -->
```python
ax = plot_nswdemand_pdf(ds)
```